Exercise 2: Advanced Gradient Descent
-------------------------------------

Before we start modifying our gradient descent algorithm, let us first
put it in a more "useful" form.  Remember from the lecture on higher-order
functions that we can pass any function `f` as parameter into another
function and then just call the passed function.

Copy your function `gradient_descent()` from the previous lecture and
modify it such that it takes an additional parameter: `grad` (which is
a function returning the gradient).

You can remove any `print()` statements from the function 
(we will do this in a better way soon).  Please also not that I renamed `alpha` 
in your previous exercise to `eta` (I messed up the notation in the theory 
part - sorry ...)

<!-- grade_id: cell-76b97c158d0e2478, points: 1, grade: true, solution: true -->
```python
def gradient_descent(grad, x0, eta, max_iter):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Now we have a truly useful gradient descent function, which works for
any 1D function we throw at it.

```python
def mexican_hat(x):
    return x**4 - 2 * x**2
    
def mexican_hat_grad(x):
    return 4 * (x**3 - x)

xmin = gradient_descent(mexican_hat_grad, -1.5, 0.1, 20)
```

```python
print("f(", xmin, ") =", mexican_hat(xmin))
```

Try this out!  Make up some potential with a minimum, define its derivative,
see if your `gradient_descent` functon can handle it.

```python

```

```python

```

```python

```

Lists and plotting
---------------------

Let's understand graphically how the `mexican_hat` function looks.  To
do this, let us first create two lists, either by appending elements in a loop or
by list comprehension:

  - a list `x` with values from `[-1.5, -1.4, ..., 1.5]`
  - a list `y` with the associated values of `mexican_hat` for each of the `x`

<!-- grade_id: cell-c03e419e7aad0621, points: 0, grade: false, solution: true -->
```python
# x = ...
# y = ...

# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-f2c0a4fc2cf428cf, points: 2, grade: true, solution: false -->
```python
assert len(x) == len(y), "Arrays x and y should be of the same length"
assert abs(x[0] + 1.5) < 1e-10, "First value of x should be -1.5"
assert abs(x[-1] - 1.5) < 1e-10, "Last value of x should be 1.5"
assert abs(y[0] - mexican_hat(-1.5)) < 1e-5, "y values do not match x"
assert abs(y[-1] - mexican_hat(1.5)) < 1e-5, "y values do not match x"

```

Let's plot this now: plot the values `y` over `x` using matplotlib. Don't forget to import the library.

Also remember: a good plot has axis labels and a title! :) 

<!-- grade_id: cell-da92b63354cbdd2b, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

Now we would like to understand how gradient descent converges to one of the minima shown above.

For this, copy the function from above and give it the new name `gradient_descent_all`.
Modify the function from its original as follows: instead of only the `x` value for the last
iteration it should give a list of values corresponding to `x` in each iteration.

<!-- grade_id: cell-968a25802446fc1a, points: 0, grade: false, solution: true -->
```python
def gradient_descent_all(grad, x0, eta, max_iter):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
gradient_descent_all(mexican_hat_grad, 1.5, 0.1, 20)
```

<!-- grade_id: cell-19bb6a0f59d85726, points: 1, grade: true, solution: false -->
```python
assert iter(gradient_descent_all(mexican_hat_grad, 1.5, 0.1, 20)), "should give list"
assert abs(gradient_descent_all(mexican_hat_grad, 1.77, 0.1, 20)[0] - 1.77) < 1e-5, "initial value missing"
assert abs(gradient_descent_all(mexican_hat_grad, 1.5, 0.1, 20)[-1] - 1) < 1e-5, "last value should converge"

```

Now you can plot this list over the iteration number `t` to see how it converges.
Make two figures, each figure with three or so plot lines: one where you vary alpha and keep the
initial position constant and one where you vary the initial position but keep
alpha constant

(Hint: look into the documentation of plot to make your life easier)

<!-- grade_id: cell-ed56b99d9b0fb926, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-84ec0aae31489e2e, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

```python

```

Nesterov-accelerated Gradient Descent
-----------------------------------------------
Finally, let's explore the effect of momentum on gradient descent.

Copy your function `gradient_descent_all`, and give it a new name
`nesterov_all`.  Modify the function to implement the Nesterov acceleration scheme.  You
now need the additional mixing parameter `gamma`.

<!-- grade_id: cell-6b7b0adc9b96bf67, points: 0, grade: false, solution: true -->
```python
def nesterov_all(grad, x0, eta, gamma, max_iter):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
nesterov_all(mexican_hat_grad, 2.0, 0.1, 0.1, 20)
```

```python

```

<!-- grade_id: cell-29fa70b257d285c9, points: 2, grade: true, solution: false -->
```python
# Sanity checks
assert iter(nesterov_all(mexican_hat_grad, 1.5, 0.1, 0.1, 20)), "should give list"
assert abs(nesterov_all(mexican_hat_grad, 1.77, 0.1, 0.1, 20)[0] - 1.77) < 1e-5, "initial value missing"
assert abs(nesterov_all(mexican_hat_grad, 1.5, 0.1, 0.1, 20)[-1] - 1) < 1e-5, "last value should converge"


```

Let's compare how Nesterov and regular GD converge: Make a figure with two plots: one where you plot the steps for regular gradient descent and one where you plot the steps for Nesterov acceleration for the `mexican_hat` function and the same starting point.

Play around a little bit

<!-- grade_id: cell-861934d7365ab5f6, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Describe in words how the behaviour changes when momentum is included.

What happens when you vary `eta`? What happens when you very `gamma`?

<!-- grade_id: cell-06be6c816ec8ba1b, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```

```python

```