Fitting Polynomials
================

In this exercise, we are going to explore the bias–variance tradeoff
on a illustrative problem, polynomial fitting.

We are now going to use the following generalized linear model
$$
f(x; \vec\theta) = \sum_{k=0}^{K-1} \theta_k f_k(x) = \theta_0 + \theta_1 x + \ldots \theta_{K-1} x^{K-1},
$$
i.e., one has $f_k(x) = x^k$.  (Note that we are indexing slightly
differently than in the lecture because numpy indices start at zero.)

Write `f_k`, which implements $f_k(x)$, and also a function `fmodel`, which implements $f(x; \vec\theta)$ above

```python
import numpy as np
import matplotlib.pyplot as pl
```

<!-- grade_id: cell-99dfa0524a14b0fe, points: 0, grade: false, solution: true -->
```python
def f_k(k, x):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def fmodel(x, theta):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-29a5d6d1a42e7b4e, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(f_k(0, 3.0), 1.0)
np.testing.assert_allclose(f_k(1, 2.5), 2.5)
np.testing.assert_allclose(f_k(2, 4.0), 16.0)

np.testing.assert_allclose(fmodel(3.0, np.array([.4])), 0.4)
np.testing.assert_allclose(fmodel(1.0, np.array([.7, 0.5])), 1.2)
np.testing.assert_allclose(fmodel(1.0, np.array([.7, 0.5])), 1.2)

```

Load and randomize data set
-------------------------------------
Instead of needing to randomly select the data by throwing dice,
a simpler technique is to just randomize the observations.  Afterwards,
we can simply select to first $n$ rows and be sure we are not biasing
our calculations.

 1. Use numpy's `loadtxt` function to again load a data matrix, where
    each row corresponds to an observation with $x$ and $y$ value,
    from a file `poly.txt`.
  
 2. Randomly permute the rows in the data matrix (Hint: look
    through the methods for `random`)
 
 3. Again split up the result into a `x` and `y` vector.

```python
random = np.random.default_rng(4712)
```

<!-- grade_id: cell-4f21642c224a9950, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print(x[:5])
print(y[:5])
```

```python

```

<!-- grade_id: cell-b34e636ead172827, points: 1, grade: true, solution: false -->
```python
assert (x, y) is not None

# Check if permutation did not mess up stuff
_sorted = np.sort(np.rec.fromarrays([x,y], names='x,y'))
np.testing.assert_allclose(_sorted.x, np.loadtxt('poly.txt')[:,0])
np.testing.assert_allclose(_sorted.y, np.loadtxt('poly.txt')[:,1])

# Check for randomness...
sortedx = (x[1:] > x[:-1]).sum()
assert sortedx > .3 * x.size and sortedx < .7 * x.size
```

Let's check if we did a good job on randomizing the data.
There is a cute way of doing this:

We again do a scatter plot (`pl.scatter`), where I plot `y`
over `x`.  These plots allow you to use **color** to attach
information to each data point.
We want color to represent the position of the observation
within the array, so we simply pass a vector with an indices.
If you like, you can play around with the [`cmap` argument].

If the resulting plot looks a little bit like an impressionist
painting, which the colors nicely scattered, then you did a
good job randomizing.  If you can spot a visible gradient, then
it means that position and index are correlated, and something
went wrong in the randomization.

[`cmap` argument]: https://matplotlib.org/stable/tutorials/colors/colormaps.html

```python
pl.scatter(x, y, c=np.arange(1000), s=5)
pl.xlabel('$x_n$')
pl.ylabel('$y_n$')
pl.title('data set')
pl.colorbar(label='$n$')
```

```python

```

Generalized linear model
------------------------------
Let's fix the maximum polynomial order at $Kmax=20$.

Now you should construct an augmented "design matrix" `Xtilde` with `N=1000`
rows, corresponding to the observations, and `Kmax=20` columns with the
element $X_{kn} = \tilde x_{n,k} = f_k(x_n)$, i.e., the $k$-th feature of the
$n$-th observation.

<!-- grade_id: cell-fa5049dd3dcf51a0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
Xtilde[:4, :4]
```

```python

```

<!-- grade_id: cell-be2ab51fbc6fe851, points: 1, grade: true, solution: false -->
```python
assert (Xtilde,) is not None
assert Xtilde.shape == (1000, 20)

_rowsums = [
    1.00000000e+03,  2.74192173e+01,  1.26907921e+03,  6.80383109e+01,
    3.03906894e+03,  1.55539066e+02,  8.77137073e+03,  3.07412343e+02,
    2.76456589e+04,  3.21997393e+02,  9.16437544e+04, -1.40425556e+03,
    3.13851105e+05, -1.36666460e+04,  1.09962081e+06, -7.87415634e+04,
    3.91828369e+06, -3.86712415e+05,  1.41452574e+07, -1.75920784e+06]
np.testing.assert_allclose(Xtilde[:,:].sum(0), _rowsums)
```

Write the corresponding model function `fmodel_lin` which maps elements `xtilde` ($\tilde x_n$)
to the predicted $y$ (note that this is now a linear model!), given the
model parameters `theta`:
$$
    \hat y_n = f_\mathrm{lin}(\tilde x_n; \vec\theta)
$$

Also implement the cost function, and include a Ridge regularization term,
weighted with `lambda_`:
$$
    \operatorname{cost}(\tilde X, \vec y; \vec\theta, \lambda) 
    =  \sum_{n=1}^N | y_n - f(\tilde x_n; \vec\theta) |^2 + \mbox{(Ridge regularization term)}
$$
You should not use hardcoded values for $N$ or $K$ since we will change them
later - infer them from the size of the arguments instead.

<!-- grade_id: cell-96a9780761ec8542, points: 0, grade: false, solution: true -->
```python
def fmodel_lin(xtilde, theta):
    # Here is a check for your convenience
    if np.ndim(xtilde) != 1 or np.shape(xtilde) != np.shape(theta):
        raise ValueError("xtilde and theta must be vectors of equal size")

    # You should implement the logic here:
    # YOUR CODE HERE
    raise NotImplementedError()
    
def cost(Xtilde, y, theta, lambda_=0):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-68408d26d2f8be09, points: 2, grade: true, solution: false -->
```python
_theta_dummy = np.exp(-np.arange(20))
np.testing.assert_allclose(fmodel_lin(np.ones(20), _theta_dummy), _theta_dummy.sum())

np.testing.assert_allclose(cost(Xtilde, y, _theta_dummy), 8088.8421161)
np.testing.assert_allclose(cost(Xtilde[:,:10], y, _theta_dummy[:10]), 8057.3235038)

np.testing.assert_allclose(cost(Xtilde, y, _theta_dummy, 1.0), 8089.99863374)
np.testing.assert_allclose(cost(Xtilde, y, _theta_dummy, 10.0), 8204.49388)
```

```python
cost(Xtilde, y, _theta_dummy, 10.0)
```

We need one final function: one which minimizes the `cost` function for given $\tilde X$, $y$
and $\lambda$ and returns the optimal parameters $\theta^*$:
$$
    \theta^* = \arg\min_\theta \operatorname{cost}(\tilde X, \vec y; \vec\theta, \lambda)
$$
Use the regularized normal equations to solve:

<!-- grade_id: cell-4ba68adb60177857, points: 0, grade: false, solution: true -->
```python
def ridge_regression(Xtilde, y, lambda_=0):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
ridge_regression(Xtilde, y)
```

<!-- grade_id: cell-e6f9d05254627411, points: 1, grade: true, solution: false -->
```python
assert ridge_regression(Xtilde, y).shape == (20,)

# Use reduced set here because of numerical instability otherwise
np.testing.assert_allclose(
    ridge_regression(Xtilde[:,:3], y), [ 1.47642395, -1.16072359,  0.69449246])
np.testing.assert_allclose(
    ridge_regression(Xtilde[:,:3], y, 1.0), [ 1.47390794, -1.15979829,  0.69529362])
np.testing.assert_allclose(
    ridge_regression(Xtilde[:,:3], y, 2.0), [ 1.46643194, -1.15703091,  0.69766407])
```

```python

```

Polynomial fit
--------------------

With all those functions we are finally ready to tackle polynomial fitting.

Again split your data, `Xtilde` and `y`, into two sets, where you should put half
into the training set (`Xtilde_train`, `y_train`) and half into the validation set
(`Xtilde_test`, `y_test`).  In total, there should be `N` observations, where
`N` can be smaller than the total number of points we have in our set.

Since we have previously randomized the data, this simplifies our job: say, N=100,
then we simply put the first 50 observations in the training set and the next 50
observations into the validation set.

We want to study how the data as we increase or reduce the polynomial order.
To do so, truncate the design matrices to the first `K` features (thereby enforcing
fitting with a polynomial of order `K`).

**Note**: numpy may get confused if you use floating point numbers as indices. You
can pass a number through the `int()` function to truncate the fractional part.

```python
K = 10
N = 100
```

<!-- grade_id: cell-839a176e820cd757, points: 0, grade: false, solution: true -->
```python
# Fill Xtilde_train, y_train, Xtilde_test, y_test here
# YOUR CODE HERE
raise NotImplementedError()
```

```python
y_test
```

```python

```

<!-- grade_id: cell-5c5e5299c4cc4287, points: 1, grade: true, solution: false -->
```python
assert (Xtilde_train, Xtilde_test, y_train, y_test) is not None
assert Xtilde_train.shape == (N//2, K)
assert Xtilde_test.shape == (N//2, K)
assert y_train.shape == (N//2,)
assert y_test.shape == (N//2,)
```

Perform Ridge regression on the training set for `lambda_=0` for now, and plot the fitted
curve as well as the points in the validation set (see last exercise).

Hint: for the plotting the $x$ values, observe that they are still present in one of the
columns of `Xtilde`.

Don't forget labels and a legend!

```python
lambda_ = 0
```

<!-- grade_id: cell-6cf95d4a3b3e0f5b, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Also compute the in error and out error and print it

<!-- grade_id: cell-375e820888cf0db0, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

print("    in-sample error: %.5g" % E_in)
print("out-of-sample error: %.5g" % E_out)
```

Execute the cells from "Polynomial fit" downwards again while varying
parameters:

 (1) while fixing `N=100`, vary the polynomial order `K`
 
 (2) while fixing `K=20`, vary the number of data points
 
 (3) take a case high `K` and low `N` and vary `lambda_`
 
What do you observe for in- and out-error and why?

<!-- grade_id: cell-4df29c3c84a7fd34, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

Analyze behaviour of Ein and Eout
--------------------------------------

Copy the above code into a function, which takes `Xtilde` and `y`
as well as `K`, `N`, and `lambda_` as inputs and returns the a
list of two elements: `E_in` and `E_out`.

<!-- grade_id: cell-efdccd6007061df8, points: 0, grade: false, solution: true -->
```python
def performance(Xtilde, y, K, N, lambda_):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
performance(Xtilde, y, 10, 100, 0.0)
```

```python

```

<!-- grade_id: cell-ae74e903b53a8846, points: 1, grade: true, solution: false -->
```python
assert len(performance(Xtilde, y, 10, 100, 1.0)) == 2

np.testing.assert_allclose(
    performance(Xtilde, y, 10, 1000, 1.0), [554.8213225, 533.8804457])
```

I will now plot the fitting peformance over the number of features
and the number of observations, respectively.

Explain these graphs below

```python
k = np.arange(1,20)
E_in, E_out = np.transpose(
    [performance(Xtilde, y, ki, 100, 0.0) for ki in k])

pl.plot(k, E_in, '-+b', label=r'$E_\mathrm{in}$')
pl.plot(k, E_out, '-xr', label=r'$E_\mathrm{out}$')
pl.title('errors vs number of features')
pl.xlabel('$k$')
pl.ylabel('$E$')
pl.yscale('log')
pl.legend()
```

```python
n = np.arange(30,200)
E_in, E_out = np.transpose(
    [performance(Xtilde, y, 10, ni, 0.0) for ni in n])

pl.plot(n, E_in, '-b', label=r'$E_\mathrm{in}$')
pl.plot(n, E_out, '-r', label=r'$E_\mathrm{out}$')
pl.title('errors vs number of observations')
pl.xlabel('$N$')
pl.ylabel('$E$')
pl.yscale('log')
pl.legend()
```

<!-- grade_id: cell-32da972b83d4a77d, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```