Dry run test assignment
-----------------------------------

You are looking at a Jupyter notebook. These notebooks will be the main form of homework
during the semester. The idea is to fill them out and also have fun with them, i.e., play
around a little bit.  For those interested/confused, there is a [tutorial] here and the complete [documentation].

Notebooks work a bit like these TI calculators (or Mathematica, if you worked with that).
You type in your code, press `Shift+Enter` or the `▸ Run` button at the top to execute it. 
Your result should then pop up.

Let's try this out: put the cursor in the next cell and execute it:

[tutorial]: https://linuxhint.com/jupyter_notebook_introduction_tutorial/
[documentation]: https://jupyter-notebook.readthedocs.io/en/stable/notebook.html

```python
1 + 1
```

Onwards! The assigments will take three main forms.  

The **first form** is a coding task, where you are supposed 
to solve a certain task in the cells that follow in whatever
way you see fit. Let's try this out:

Use the following cells to do a couple more simple 
calculations (you can use `+`, `-`, `*`, `/`, etc.) and
execute them.  Add a few more cells by using the `+` button in the toolbar above or by using `Insert > Insert Cell below`

```python

```

```python

```

```python

```

The **second form** is a bit like a "Lückentext" (cloze):
you get a cell that is already filled, but then there is some bits missing and the marker `raise NotImplementedError()`   You are supposed to **replace** that marker with your code like in the assignment.

Again let's try this:  delete the `raise ...` and instead put the statement `x = 1` below:

<!-- grade_id: cell-b6a4c8d5cb64c8ac, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

x + 5
```

Especially for the second form, we will ofter provide you with "test cells", which check whether your solution is correct.

If you filled the above cell correctly, then executing the following cell should work.  Otherwise you will get an error:
see what happens if you change your above answer to `x = 2`, execute both cells again, then change it back.

<!-- grade_id: cell-40868e5f1db1ff21, points: 1, grade: true, solution: false -->
```python
assert x == 1
```

```python

```

The **third form** is just text answers, where you supposed to
give details to the method, discuss results, etc.

They live in special cells which contain text rather than code.  You recognize text cells by the fact that they don't have the `In [ ]:` at the beginning and the test `YOUR ANSWER HERE`.

Let's try this form as well: **double click** the text `YOUR ANSWER HERE` and write some more text in it.  You could, e.g., describe whether the instructions here were easy to follow or not.

<!-- grade_id: cell-77cd5c633643a69e, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```

We're done.

Please hand in your notebook.

 1. Save this file `Ctrl+S` or `File > Save and checkpoint`
 2. Go to the main page of Jupyter
 3. Go to the assignments tab
 4. Recommended: Click the "Validate" button on the assigment 
 5. Hand it in. All done!

```python

```