Exercise 3: Newton's method
===========================

We are going to implement Newton's method in multiple dimensions.

We are going to consider the following function:
$$
        E(x) = ||x||^4 - ||x||^2 - \frac 15 x_0,
$$
where $x = (x_0, x_1)^T$ and $||x||^2 = x_0^2 + x_1^2$.
Using $\partial_i ||x||^2 = 2 x_i$, one can work out the gradient
of this function:
$$
        J_i(x) = \frac{\partial E(x)}{\partial x_i} = (4 ||x||^2 - 2) x_i - \frac 15 \delta_{i0}
$$
First, write a function $E(x)$ which takes a numpy array $x$ as input and computes the energy
at $x$ and a function $J(x)$ which returns a numpy vector for the gradient at $x$.

Don't forget to import the numpy package.

<!-- grade_id: cell-3788c1aeffb978c0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

def E(x):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def J(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
E(np.array([-0.4,0.1]))
```

```python
J(np.array([0.1,-0.2]))
```

<!-- grade_id: cell-c4a4f0ff742bd35a, points: 2, grade: true, solution: false -->
```python
assert np.allclose(E(np.array([0.0,0.0])), 0, atol=1e-4)
assert np.allclose(E(np.array([0.5,0.2])), -0.3059, atol=1e-4)
assert np.shape(J(np.array([0.0,0.0]))) == (2,)
assert np.allclose(J(np.array([0.752619,0.0])), [0,0], atol=1e-4)

```

For your convenience, there is a false color plot of $E(x)$ below:

```python
# You don't need to understand this code
def _plot(f):
    import matplotlib.pyplot as pl
    x = np.linspace(-1.2, 1.2, 31)
    fx = np.array([[f(np.array([a, b])) for a in x] for b in x])
    pl.pcolormesh(x, x, fx, cmap='seismic', vmax=1, vmin=-1, shading='auto')
    pl.colorbar()
    pl.xlabel(r'$x_0$')
    pl.ylabel(r'$x_1$')
    pl.title(r'$E(x)$ with $x = (x_0, x_1)^T$')
    pl.axis('scaled')
    x0 = [-0.6505, -0.10213, 0.75262]
    pl.scatter(x0, [0, 0, 0], marker='x', c='k')
    for x0i, t in zip(x0, 'ABC'): pl.text(x0i, 0.1, t, ha='center')

_plot(E)
```

In the above plot, I have marked points (A,B,C) where the gradient is zero:

$$
\begin{aligned}
 x_A &\approx (-0.6505, 0)^T \\
 x_B &\approx (-0.10213, 0)^T \\
 x_C &\approx (0.75262, 0)^T \\
\end{aligned}
$$

Identify the character of these points below:

<!-- grade_id: cell-cb2974020bc029fc, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Adapt your gradient descent solver (without Nesterov acceleration) from the
previous exercise so that it can handle any dimension.

`x0` is then a numpy array corresponding to the first points, `grad` is a
function taking and returning an array.  Be mindful: in numpy, comparisons are
component-by-component (this may lead to an error).

<!-- grade_id: cell-39a3f30255cf24b4, points: 1, grade: true, solution: true -->
```python
def gradient_descent(grad, x0, eta, max_iter):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
gradient_descent(J, np.array([0.1, 0.1]), 0.1, 20)
```

```python

```

Note that the minimum found by the above code is off by quite a bit ... why?

How can you change the parameters so that it works better?  Would acceleration help?

<!-- grade_id: cell-e91fcd97332874ec, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python
# Cells for trying out parameters ...
```

```python

```

```python

```

```python

```

```python

```

Hessian
------------

The Hessian of $E$ is given by the second derivatives, which are:
$$
    H_{ij}(x) = \frac{\partial^2E(x)}{\partial x_i\partial x_j}  = (4 ||x||^2 - 2) \delta_{ij} + 8 x_i x_j
$$
Implement this in the function $H$ below, which returns a $2\times 2$ numpy matrix.

<!-- grade_id: cell-b7dbd95fa1e5009d, points: 0, grade: false, solution: true -->
```python
def H(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
H(np.array([0.1,0.3]))
```

<!-- grade_id: cell-4a7c8171d51d2b5f, points: 1, grade: true, solution: false -->
```python
assert H(np.array([0.0,0.0])).shape == (2, 2)
assert np.allclose(H(np.array([0.0,0.0])), [[-2, 0], [0, -2]], atol=1e-5)
assert np.allclose(H(np.array([1.0,2.0])), [[26, 16], [16, 50]], atol=1e-5)

```

Let's remind ourselves of the points (A,B,C) where the gradient is zero:

$$\begin{aligned}
 x_A &\approx (-0.6505, 0)^T \\
 x_B &\approx (-0.10213, 0)^T \\
 x_C &\approx (0.75262, 0)^T \\
\end{aligned}$$

Compute the Hessian for these points and prove the character you identified above
using the Hessian.

```python

```

```python

```

```python

```

```python

```

```python

```

Summarize your findings below:

<!-- grade_id: cell-36b2f7c41442e9ca, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Newton's method
---------------
Let's combine everything and implement Newton's method.
You can start, e.g., from the code of your `gradient_descent` method.

Newton's method shall take a function for the gradient `grad` a function for
the Hessian `hess`, a start point `x0`, a small non-negative regularization parameter `lambda_`
(note the trailing underscore), a learning rate `eta` between 0 and 1, and
the maximum  number of iterations.

Hint `x.size`  gives the total number of elements in an array.

<!-- grade_id: cell-574a9c4248f960f5, points: 0, grade: false, solution: true -->
```python
def newton(grad, hess, x0, lambda_=0, eta=1, max_iter=20):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
newton(J, H, np.array([0.5, 0.0]))
```

<!-- grade_id: cell-ad5f11b7187b384f, points: 2, grade: true, solution: false -->
```python
assert newton(J, H, np.array([0.0, 0.0])).shape == (2,)
assert np.allclose(newton(J, H, np.array([0.0, 0.0]), max_iter=0), [0.0, 0.0], atol=1e-4)
assert np.allclose(newton(J, H, np.array([0.8, -0.1])), [0.752619, 0.0], atol=1e-4)
assert np.allclose(newton(J, H, np.array([0.4, 0.1])), [0.752619, 0.0], atol=1e-4)

```

Compare the results for Newton's method and Gradient Descent for the starting point $x = (-0.5, 0.05)$. (For gradient descent you probably need a large number of iterations.)

Also investigate the result for the starting point $x = (0, 0)$.

```python

```

```python

```

```python

```

```python

```

```python

```

What happens in each approach and why?

<!-- grade_id: cell-0d2ff07b1389ef52, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```