Exercise 10: Phase transitions in the Ising Model
=================================================
<img src="ising.png" style="max-width:30%; float:right; padding-left:30pt">

The [Ising model] is perhaps the most important model in all of solid state physics, certainly
one of the most significant from a historical perspective.  This is because it is a simple
microscopic, atomistic model from which a (macroscopic) magnetic phase transition can be shown to *emerge*,
finally laying to rest a debate over whether statistical mechanics works and even whether atoms are real that continued well into the twentieth century.

The idea is the following: consider a $N\times N$ square grid of atoms (we will take $N = 40$).  The atom at each position $(i,j)$ is fixed in place, but is endowed with a **spin** $\sigma_{i,j}$ which can either be $+1$ or $-1$ (pointing up or down) at any given moment.  One such configuration can be seen on the right: the top-left spin $\sigma_{1,1}=1$, the one two columns to the right $\sigma_{1,3}= -1$, and so forth.  The potential energy $U$ encoded in each configuration is given by:

$$
   U = -\sum_{i,j} \big[ \sigma_{i,j} \sigma_{i,j+1} + \sigma_{i,j} \sigma_{i+1,j} \big],
$$

in other words, we have a positive contribution to the potential energy $\Delta U = 1$ whenever two neighbouring  spins are opposite to each other and a negative energy contribution $\Delta U = -1$ whenever two neighbouring spins are aligned.  Competing with the systems desire to lower its potential energy is temperature $T$, which introduces kinetic energy by randomly flipping spins.

We want to explore this competiton between kinetic and potential energy.  For this, I  have prepared a set of $N$ snapshots of the spins $\sigma^{(n)}_{i,j}$ taken at different times, stored in the $N\times L \times L$ tensor `spins`, together with their respective temperatures $T_n$, stored in a vector `temp`.

Each of these observations may belong to a "phase" $y_n$ of the material, which is unknown to us.  However, we suspect that similar phases have similar properties.  The idea is find different these without any knowledge of the phases themselves by using k-means clustering.

[Ising model]: https://en.wikipedia.org/wiki/Ising_model

```python
import os
import numpy as np
import matplotlib.pyplot as pl
```

```python
# Load the dataset from a binary file
with np.load("../shared/ising.npz") as _datafile:
    spins = _datafile["spins"]
    temp = _datafile["temp"]
```

```python

```

Let us first make a warm-up plot: make a false-color plot of the $n=140,000$'th spin configuration
similar to the figure above, where you use one color to denote spin-down ($\sigma_{ij}=-1$)
and another color to denote spin-up ($\sigma_{ij}=1$).

<!-- grade_id: cell-46f5abb99e04f564, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

```python

```

Step 1: Compute observables
---------------------------
Clustering in high dimensions (which is what we would do if we work with the spin array directly) is
difficult to perform and also visualize.  Therefore we are going to reduce each configuration ($s^{(n)}_{i,j}$
for a fixed $n$) to just two observables: the **magnetization** $m$, which we will store in `magn`, and the
**staggered magnetization** $m'$, which we will store in `smagn`.

The definitions are:
$$
    m_n = \frac 1{L^2} \sum_{i,j} \sigma^{(n)}_{i,j}
$$
$$
    m'_n = \frac 1{L^2} \sum_{i,j} (-1)^{i+j} \sigma^{(n)}_{i,j}
$$

Compute the two numpy vectors `magn` and `smagn` according to above formula

<!-- grade_id: cell-e7eb900b0043e8f8, points: 0, grade: false, solution: true -->
```python
#magn = ???
#smagn = ???

# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-f924b10e094d4a1b, points: 2, grade: true, solution: false -->
```python
assert (magn, smagn) is not None
assert magn.shape == smagn.shape == (160_000,)
assert np.all(np.abs(magn) <= 1)
assert np.all(np.abs(smagn) <= 1)

np.testing.assert_allclose(magn.sum(), 191.63875, rtol=1e-4)
np.testing.assert_allclose(smagn.sum(), 5.68875, rtol=1e-4)
```

Now, perform a scatter plot, where each point is one observation, the $x$-axis corresponds to the magnetization of that observation, the $y$ axis corresponds to the staggered magnetization of the observation and the color corresponds to the temperature.  Also add a colorbar

Choose a small marker size (`s` argument), otherwise you not see much in the plot

<!-- grade_id: cell-7e31cd8358383684, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

Step 2:  Clustering, first round
--------------------------------
Your plot above should suggest to use three clusters.

Perform a cluster analysis using the `sklearn.cluster.KMeans` clustering
algorithm.  For this, first collect `magn` and `smagn` as columns of a `N x 2` design matrix
`X` and then run the fitting.

```python
import sklearn.cluster
```

<!-- grade_id: cell-b16c9992f865207a, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

Now, repeat the scatter plot from above, but instead of temperature, plot each point together in
different colors depending on their cluster label.

<!-- grade_id: cell-c171188788ca2568, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

Let us get a feeling for the kind of configurations in each of the clusters:
for each cluster, pick two configurations at random and make a false color
plot for them.  You should therefore end up with a 3 x 2 matrix of plots,
where the columns correspond to the different clusters and rows are the
different configurations

**Note**: if all the spins are +1 or -1, then the colorbar gets confused,
since there is only one value.  You can pass  `(..., vmin=-1, vmax=1)` arguments
to false color plots to combat this.

<!-- grade_id: cell-f40330200e55e596, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-d2d38214803bdab4, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

YOUR ANSWER HERE

```python

```

Step 3: Clustering, second round
--------------------------------
We would like to take the symmetry of our system into account.  Therefore,
instead of the (staggered) magnetization, we are using the **absolute**
values of each of these quantities.

Replace `magn` and `smagn` with a vector of their absolute values and
redo the scatter plot over temperature.

<!-- grade_id: cell-4dd2d0739a26b92b, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

Now, redo the cluster analysis for these new variables and plot the result.

Again use **three** clusters.

<!-- grade_id: cell-f0b4b8e3ec5c7b3f, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

Repeat the grid plot where you pick two configurations for
each cluster:

<!-- grade_id: cell-6cd583ad35042d60, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-4d17b9c733cad0db, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

YOUR ANSWER HERE

```python

```