```python
# <-- This marks a comment until the end of the line
import math
```

Gradient Descent in 1D
---------------------------------
Please first watch the videos for Python and Gradient Descent.

We are going to perform Gradient Descent (GD) for the simple function:
$$
    f(x) = x^4 - 2 x^2
$$

Complete the function below and try a few values:

<!-- grade_id: cell-b201df9927df4ef7, points: 0, grade: false, solution: true -->
```python
def f(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
f(0)
```

```python

```

```python

```

```python

```

<!-- grade_id: cell-296fee55e3fa2506, points: 1, grade: true, solution: false -->
```python
assert f(0) == 0
assert f(1) == -1
assert f(-1) == -1

```

For Gradient Descent, we need the derivative.

Write a function names `dfdx`, which takes a parameter `x` and
returns $df/dx$ at that position:

<!-- grade_id: cell-ea6ad626eb794300, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
dfdx(0.2)
```

<!-- grade_id: cell-221d948f68c10aea, points: 1, grade: true, solution: false -->
```python
assert abs(dfdx(0.5) + 1.5) < 1e-5

```

Here are some initial values for Gradient descent.

In the cell below, perform a **single** step of Gradient descent,
i.e., given a value of `x`, update it according to the slope and
learning rate `eta`. Print the value of `x` at the end of the
cell.

```python
eta = 0.1
x = -1.5
```

<!-- grade_id: cell-31d499e5ce33e8e4, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
print("x = ", x)
```

Now execute the above cell multiple times ... observe how x is updated.

Let's have Python take over the task of performing multiple steps. Write
a loop below that updates `x` according to the Gradient Descent rules
until one of two conditions are met:

 1. the absolute value of the gradient is less than $10^{-5}$ 
 2. `max_iter` iterations have been performed
 
After each iteration, print the value of `x`.

```python
max_iter = 10
eta = 0.1
x = -1.5
```

<!-- grade_id: cell-166939ca3308ae0f, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

For the coup de grâce, let's wrap the whole thing in a
function.

Write a function `gradient_descent`, which takes three
arguments:

 1. the initial value `x`, 
 2. the learning rate `eta`
 3. and the maximum number of iterations `max_iter`.
 
It should then update the value `x` according to Gradient
descent (see loop above) and finally return the approximation to
the minimum.  

<!-- grade_id: cell-d55fff0b84243184, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
gradient_descent(-1.5, 0.1, 20)
```

<!-- grade_id: cell-cea517014b379ef2, points: 2, grade: true, solution: false -->
```python
np.testing.assert_allclose(gradient_descent(-1.5, 0.1, 20), -1, atol=1e-5)
np.testing.assert_allclose(gradient_descent(0.8, 0.15, 20), 1, atol=1e-5)
np.testing.assert_allclose(gradient_descent(0.8, 0.1, 0), 0.8, atol=1e-5)

```

You should now play around with your function:

 1. try different starting points. What happens when starting at $x = 0$ or $x$ close to a minimum?
 2. try different values of $\eta$.  How does the number of needed iterations depend on that?
 3. what happens if you make $\eta$ large?
 4. can you find an $\eta$ such that the minimum can be reached in a single step?

You can record your answers to questions in the text block further below:

```python
# Cells for playing around
```

```python

```

```python

```

```python

```

<!-- grade_id: cell-56ae11f2f00aec1c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```