Exercise 6: singular value decomposition
========================================

In this exercise we are going to explore the singular value decomposition.
To start things off, let us flex our linear algebra muscles a little bit.

Let
$$
     X = \sum_{k=0}^{K-1} s_k \vec u_k \vec v_k^T
$$
be the singular value decomposition of $X$.
Prove that the SVD of the regularized pseudoinverse of $X$ can be written as:
$$
    (X^T X + \lambda^2 \mathbf 1)^{-1} X^T = \sum_{k=0}^{K-1} \frac{s_k}{s_k^2 + \lambda^2} \vec v_k \vec u_k^T
$$
Do we need to assume $s_k \neq 0$ in the derivation? Why or why not?


<!-- grade_id: cell-ee5c87ef27c52927, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```

Polynomial fitting and the Vandermonde matrix
---------------------------------------------

Let us explore polynomial fitting a little more, since it is a very instructive
example for SVD and regularization.  Remember that given some $x_n$ and labels $y_n$
we wanted to fit a model:
$$
    \arg\min_\theta \sum_{n=0}^{N-1} \Big| y_n - \sum_{k=0}^{K-1} \theta_k x_n^k \Big|^2
    = \arg\min_\theta || \vec y - X \vec\theta ||^2
$$

The design matrix in above equation is called a [**Vandermonde matrix**](https://en.wikipedia.org/wiki/Vandermonde_matrix):
$$
    X_{nk} = x_n^{k}
$$
and it is famously nasty if you're not careful, which is why you might have had problems fitting
in the previous exercise.

Let us explore this is more detail by constructing our own fitting problem: construct a vector `x` of
`N = 101` values linearly spaced in the interval `[-1, 1]`. The construct the Vandermonde matrix `X`
for `K = 30` (i.e., maximum polynomial order 29).

<!-- grade_id: cell-260e411e1e3424e2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e98d58657bc2e1a1, points: 1, grade: true, solution: false -->
```python
import numpy as np

assert x.shape == (101,)
np.testing.assert_allclose(np.diff(x), 0.02)

assert X.shape == (101, 30)
np.testing.assert_allclose(np.linalg.cond(X), 51785875457, rtol=1e-3)
```

```python

```

Let us analyze the design matrix: 

 1. compute the SVD of `X` using numpy's linear algebra routines.
    Store the resulting left and right singular matrices in `U` and `VT`, respectively,
    and the vector of singular values in `s`.
   
 2. Afterwards, reconstruct the original matrix from the SVD factors and store it into `X_rec`. 

**Note**: In the lecture I presented a variant of the SVD called "thin SVD", where `U` and
`VT` are not necessarily square.  Please set the `full_matrix` argument to do so. 

<!-- grade_id: cell-ff9809291ee4a64c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8d7ea9e4e96bee63, points: 1, grade: true, solution: false -->
```python
assert (U, s, VT, X_rec) is not None
assert s.shape == (30,)
assert U.shape == (101, 30)
assert VT.shape == (30, 30)
np.testing.assert_allclose(X_rec, X, atol=1e-5)
```

Finally, write a function `pinv()` that takes a matrix in a singular-value decomposed
from (`U`, `s`, `VT`) and constructs the pseudo-inverse:

<!-- grade_id: cell-7732e3b9c943ce1c, points: 0, grade: false, solution: true -->
```python
def pinv(U, s, VT):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
pinv(U, s, VT)
```

<!-- grade_id: cell-16540cdf1dbcb8a4, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(pinv(U, s, VT), np.linalg.pinv(X))
```

Analyzing the singular value decomposition
------------------------------------------
Let's get a little bit of intuition on the factors in
the SVD

First, plot the singular values on a logarithmic scale.

<!-- grade_id: cell-66aa64f5b421e8f1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Next, make a figure where you plot the a elements of a
single left singular vector $\vec u_k$.  You should
make three plots in total, one for $k=0$, $k=6$
and $k=K-1$, respectively.

Think about the $x$-axis: what does it represent (Hint: how
did you construct $X$)? Choose the axis scaling and the axis label
accordingly.  Don't forget plot labels and a legend.

<!-- grade_id: cell-f609f0dbcc13639c, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Let's think about what will happen if we do fitting with
a Vandermonde design matrix `X`.  Remember that fitting a
set of labels `y` is equivalent to multiplying `y` with the
pseudoinverse of `X` (what is its SVD?)

  1. First, think about the singular values. What happens 
     in the fit for $k\approx 0$, what for $k\approx K$?
     
  2. Imagine $y$ to be some polynomial.  What is the a
     connection between polynomial order and left singular
     vectors?  Why do you think that is?
  
  3. Now imagine $y$ to be noisy and look at the shape of
     the left singular vector.  Which singular vectors do
     you expect to pick up most of the noise?  Why?
     
Observe that we can do this analysis *purely* on the design
matrix without looking at the labels vector, i.e., we uncover
intrinsic and fundamental features of fitting polynomials.

<!-- grade_id: cell-3e8c215238211930, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```

For completeness, also plot the right singular vectors $\vec v_k$.
You should again make three plots in total, one for $k=0$, $k=6$
and $k=K-1$, respectively.

Again think of what the axis represents.

Note that in `VT` the singular vectors are the rows, not the columns.

<!-- grade_id: cell-29a967c638cca045, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

Generating truth and noise
--------------------------
Let us try our luck by fitting a fourth-order polyomial, which for
the moment is our "truth".  Below I have defined $\theta_\mathrm{true}$, which
are the true parameter set $\vec\theta_\mathrm{true}$ and the corresponding model
$y_\mathrm{true}(x) = f(x, \vec\theta_\mathrm{true})$.

```python
# The Truth (TM)
theta_true = np.zeros(30)
theta_true[:5] = [-0.05, .13, .87, 0, -1]

def get_y_true(x):
    """True labels (without noise) for a given `x` in `[-1, 1]`"""
    x = np.asarray(x)
    if not np.all(x >= -1) and np.all(x < 1):
        raise ValueError("x must be between -1 and 1")
    return (x[..., np.newaxis]**np.arange(30) * theta_true).sum(-1)
```

Let us now **generate some labels** from this truth plus some noise.

Generate the vector of true labels `y_true` for your $x$ values. Then,
generate a "noise vector" `epsilon` where you draw each element from a 
normal distribution with mean `0` and standard deviation of `0.1`.
Finally `y_noisy` should be the sum of truth plus noise.

**Plot** the true data and the noisy data over `x`.

```python
rng = np.random.default_rng(4711)
```

<!-- grade_id: cell-7f874d9029eb76e2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f743c64ceb3ec555, points: 1, grade: true, solution: false -->
```python
assert (y_true, epsilon, y_noisy) is not None
np.testing.assert_allclose(y_true, y_noisy - epsilon)

_stddev = np.sqrt(np.linalg.norm(y_noisy - y_true)/y_true.size)
np.testing.assert_allclose(_stddev, 0.1, atol=0.03,
                           err_msg="noise not distributed correctly")
```

```python

```

Fitting data
------------
Let us now try to fit our noisy labels `y_noisy`.

To do so, use `pinv()` to compute the pseudoinverse of `X`
and then use that to compute `theta_star`.  Plot your
fitting parameters together with the true parameters
`theta_true`.

<!-- grade_id: cell-5ed584b6a43e8829, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

 1. Observe the scale of the y axis in the above plot. (The noisy data is clearly
    several orders of magnitudes "too big" with respect to the truth). Why?
    What is the connection of the scale of the coefficient with the singular values?
    
 2. The shape of the noisy coefficients should look familiar (check with the 
    plots in the singular value decomposition).  Explain this.

YOUR ANSWER HERE

```python

```

Regularization
--------------
Finally, let us regularize our problem.  One of the simplest regularization
techniques is to just discard use Ridge regression.  Again, we
choose:
$$
    X^\oplus_\lambda = \sum_{k=0}^{K-1} \frac{s_k}{s_k^2 + \lambda^2} \vec v_k \vec u_k^T
$$

Copy the function `pinv` into `pinv_ridge` and modify is such that it computes the
above expression

<!-- grade_id: cell-8fdd73a4579a296d, points: 0, grade: false, solution: true -->
```python
def pinv_ridge(U, s, VT, lambda_):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
pinv_ridge(U, s, VT, 0.1)
```

<!-- grade_id: cell-c86c6e3bd554b731, points: 1, grade: true, solution: false -->
```python
_reg_pinv = np.linalg.pinv(np.vstack([X, 0.1 * np.eye(30)]))[:, :101]
np.testing.assert_allclose(pinv_ridge(U, s, VT, 0.1), _reg_pinv, atol=1e-5)
```

```python

```

Plot $\theta^* = X^\oplus y$ for three different values of $\lambda$:
1, 0.1, and 0.01.  Also plot the "true" parameters for reference.

<!-- grade_id: cell-720198798143cb09, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

Finally, plot the predicted values $\hat y = X \theta^*$ for
the parameters fitted with $\lambda$: 1, 0.1, and 0.01.  Again,
plot the true $y$ for reference.

<!-- grade_id: cell-958fe0dc563520e0, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Summarize your results:

 1. which value of lambda gives the best fit? How can you explain this from the error and the SVD?
 
 2. why are the parameters `theta` for small `lambda` changing so much more than the predicted label `y`?
 What is going on there?

<!-- grade_id: cell-8b924c369eb8c50e, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```